---
title       : Creating your First Neural network
description : Create your first Neural Net using neuralnet package and train to predict the squares





--- type:NormalExercise lang:r xp:100 skills:1 key:4a8187545c
## Excercise 1: Importing neuralnet package


*** =instructions
To start let's import the package neuralnet which has the required libraries to finish our first tour with the neuralnet.
*** =hint
import neuralnet package using  library()
*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}
library(neuralnet)

```

*** =sct
```{r}

```
---
title       : Creating your First Neural network
description : Create your first Neural Net using neuralnet package and train to predict the squares

--- 
